---
rule_files:
  - ${REPO_ROOT}/prometheus/rules/quality_slo.rules.yml
  - ${REPO_ROOT}/prometheus/rules/slo_constants.rules.yml
evaluation_interval: 1m

tests:
# 위반 직전 경계값 테스트 (알람이 안 뜨는지 확인)
- interval: 1m
  input_series:
    - series: 'duri_mrr_at_k{k="3",scope="all",domain="ALL"}'
      # 정상(0.90) 30분 → 임계 직전(0.881) 20분 → eval_time 50m
      # MA7 계산: (0.90*30 + 0.881*20) / 50 = 0.8925490196078432 > 0.88 (SLO target)
      values: '0.90x30 0.881x20'
  promql_expr_test:
    - expr: duri_mrr_ma7
      eval_time: 50m
      exp_samples:
        - labels: '{__name__="duri_mrr_ma7", domain="ALL", k="3", scope="all"}'
          value: 0.8925490196078432

# 회복 직후 경계값 테스트 (알람이 안 뜨는지 확인)
- interval: 1m
  input_series:
    - series: 'duri_mrr_at_k{k="3",scope="all",domain="ALL"}'
      # 위반(0.86) 20분 → 회복(0.90) 30분 → eval_time 50m
      # MA7 계산: (0.86*20 + 0.90*30) / 50 = 0.8835294117647059 > 0.88 (SLO target)
      values: '0.86x20 0.90x30'
  promql_expr_test:
    - expr: duri_mrr_ma7
      eval_time: 50m
      exp_samples:
        - labels: '{__name__="duri_mrr_ma7", domain="ALL", k="3", scope="all"}'
          value: 0.8835294117647059
