rule_files:
  - ${REPO_ROOT}/prometheus/rules/quality_slo.rules.yml
evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
    - series: 'duri_mrr_at_k{k="3",scope="all",domain="ALL"}'
      # 7일치 데이터: 처음 6일은 정상(0.90), 마지막 1일은 임계 하회(0.86)
      # MA7 계산: (0.90*6 + 0.86*1) / 7 = 0.895 < 0.88 (SLO target)
      values: '0.90x6 0.86x1'
  promql_expr_test:
    - expr: duri_mrr_ma7
      eval_time: 7m
      exp_samples:
        - labels: '{__name__="duri_mrr_ma7", domain="ALL", k="3", scope="all"}'
          value: 0.895
