rule_files:
  - simple_alerts.rules.yml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: 'duri_ndcg_at_k{k="3",scope="all",domain="ALL"}'
    values: '0.90 0.88 0.79 0.78 0.77 0.79'

  alert_rule_test:
  - eval_time: 6m
    alertname: QualityDropImmediate
    exp_alerts:
      - exp_labels: {severity: "critical", team: "search"}
